<a href="https://colab.research.google.com/github/dannyxia7/UCSD_DSC_207R/blob/main/Danny_Xia_Week1_Assignment_Part_B.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Week 1 Study Notebook

# UC San Diego OMDS DSC 207

## Check your Python skills while using Colab notebooks.

### Installing Otter-Grader

Otter-Grader is a powerful tool that helps you check your work before submitting assignments. To get started with Otter-Grader, you need to install it within this Colab notebook.

To install Otter-Grader, simply run the below code cell.

Note that questions 1-9 are autograded, while question 10 will be graded manually.

```python
# DO NOT MODIFY
%pip install otter-grader

files = "https://github.com/dsc207rfall2023/fa23course-dsc207r/raw/main/assignments/Week%201/tests.zip"
!wget $files && unzip -o tests.zip
```

Output:
```
Collecting otter-grader
  Downloading otter_grader-5.7.1-py3-none-any.whl.metadata (2.9 kB)
Collecting dill (from otter-grader)
  Downloading dill-0.3.9-py3-none-any.whl.metadata (10 kB)
Requirement already satisfied: jinja2 in /usr/local/lib/python3.10/dist-packages (from otter-grader) (3.1.4)
Requirement already satisfied: nbformat in /usr/local/lib/python3.10/dist-packages (from otter-grader) (5.10.4)
Requirement already satisfied: pandas in /usr/local/lib/python3.10/dist-packages (from otter-grader) (2.2.2)
Requirement already satisfied: PyYAML in /usr/local/lib/python3.10/dist-packages (from otter-grader) (6.0.2)
Collecting python-on-whales (from otter-grader)
  Downloading python_on_whales-0.73.0-py3-none-any.whl.metadata (18 kB)
Requirement already satisfied: requests in /usr/local/lib/python3.10/dist-packages (from otter-grader) (2.32.3)
Requirement already satisfied: wrapt in /usr/local/lib/python3.10/dist-packages (from otter-grader) (1.16.0)
Collecting jupytext (from otter-g
```

```python
# Initialize Otter
import otter
grader = otter.Notebook()
```

### Loading Python modules

```python
import pandas as pd
import numpy as np

# ADD ADDITIONAL PYTHON MODULES BELOW
import math
```

### 1.  Write a program to swap a string and an int variable and print their data types to verify the swapping operation

* Initialize 2 variables my_int = 5 , and my_str = "Hi"
* Swap these variables such that my_int now holds the value of my_float and vice versa
* Verify this by printing their data types

_Points:_ 5

```python
my_int = 5
my_str = "Hi"
dummy_var = my_int
my_int = my_str
my_str = dummy_var
type(my_int)
type(my_str)
```

Output:
```
int
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

### 2. Write a Python Program to Calculate the length of the diagonal of a rectangle

* Initialize the 2 sides of the rectangle a and b to 5 and 6
* Calculate the length of the diagonal using the formula - diagonal = √(a^2 + b^2)

_Points:_ 5

```python
def calc_diag(a, b):
    return math.sqrt(a**2 + b**2)
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

### 3. Write a Python Program to Check if a Number is divisible by 3

* Initialize a variable that takes input from user
* Check and return True or False if the number is divisible by 3
* If the input is not an int, return False

_Points:_ 10

```python
def is_divisible_by_3(num):
    if type(num) != int:
      return False
    if num % 3 == 0:
      return True
    else:
      return False
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

<!-- BEGIN QUESTION -->

### 4. Write a Python Program to Check if a number has exactly 3 digits (no zeros allowed at the beginning and the number can be assumed to be positive)

* Initialize a variable that takes input from user(input will be given as a string)
* Check and print if the number has exactly 3 digits(no zeros allowed at the beginning)
* Read about type casting in python and use type casting from string to int

_Points:_ 10

```python
def is_3_digits():
  num = input("type in a number: ")
  if num[0] == "0":
    print("number can't start with 0")
  num = int(num)
  if num % 1000 == 0:
    print("nunber has exactly 3 digits")

is_3_digits()
```

Output:
```
type in a number: 040
number can't start with 0
```

<!-- END QUESTION -->

### 5. Write a Python Program to Check Whether a String is Palindrome or Not

* Initialize a string
* Lower case it to allow easy comparison
* Check using if-else statements if the string is a palindrome
* A palindrome is a string which is same read forward or backwards.

_Points:_ 10

```python
def if_palindrome(x):
    x_str = str(x).lower()
    i = 0
    j = len(x_str) - 1
    while i < j:
      if x_str[i] != x_str[j]:
        return False
      i += 1
      j -= 1

    return True
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

### 6.  **Print numbers divisible by 5 between 4 and 30, inclusive**

* Iterate through the numbers in the for loop and notice the way range has been used to generate numbers in the for loop
* Produce a list and print containing only the numbers divisible by 5 using if statement

_Points:_ 5

```python
sol = []
for i in range(4, 31):
  if (i % 5) == 0:
    sol.append(i)
print(sol)
```

Output:
```
[5, 10, 15, 20, 25, 30]
```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

### 7. Write a Python Program to print the first n numbers of the Fibonacci series starting from 0

* Use a loop and return a list of first n numbers of the Fibonacci series: 0,1,1,2,3,....


_Points:_ 15

```python
def fibonacci(n):
  if n == 0:
    return []
  elif n == 1:
    return [0]

  fib_list = [0, 1]
  for i in range(n-2):
    fib_list.append(fib_list[-1] + fib_list[-2])
  return fib_list
```

Output:
```
[0, 1, 1, 2, 3]
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

### 8. Write a Python Program that defines a function to Remove Vowels From a String

* You can initialize a string of your choice
* Define the function
* Use a for loop to return an updated string with the vowels removed.

_Points:_ 10

```python
def remove_vowels(input_str):
  vowels = ['A', 'E', 'I', 'O', 'U', 'a', 'e', 'i', 'o', 'u']
  new_str = ""
  for i in range(len(input_str)):
    if input_str[i] not in vowels:
      new_str += input_str[i]
  return new_str


```

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
```

### 9. Write a program to define a function that takes two numbers as arguments and returns the sum of all numbers between the arguments passed (exclusive, i.e. excluding the limits).

* You can define a range of your choice
* Define the function
* Use a for loop to add numbers within the range(exclusive)

_Points:_ 10

```python
def sum(a, b):
  range_sum = 0
  for i in range(a+1, b):
    range_sum += i
  return range_sum

print(sum(1,5))
```

Output:
```
9
```

```python
grader.check("q9")
```

Output:
```
q9 results: All test cases passed!
```

<!-- BEGIN QUESTION -->

### 10. Define your own Python function!

#### Part A
In Python, functions are blocks of code that only run when they are called. Functions are programmed to complete a specific task. You can input data, called parameters, into functions. Functions can also return data as a result.

Write a function with the following components:

* Pass a list as the function parameter. The list can be composed of any data type (int, string, float, etc.)
* At least one conditional (if/else/elif)
* One loop (ex. for, while, do-while)
* Returns one or more values. This value can be a list or a single element.


Ensure the following:
* Your program contains all the components listed above.
* You test the function by calling it outside of the function.
* Your program runs without any errors.

_Points:_ 10

```python
#Find numbers that add up to target in the list. Exactly one answer.
def twoSum(nums, target):
  h_map = {}
  for i in range(len(nums)):
    if target - nums[i] in h_map:
      return [h_map[target - nums[i]], i]
    else:
      h_map[nums[i]] = i

print(twoSum([1,4,12,25], target = 13))
print(twoSum([3, 7, 42, 94], target = 101))




```

Output:
```
[0, 2]
[1, 3]
```

#### Part B
Record a video, no more than 3 minutes long.

* Explain the function at a high level. We want you to explain the flow of code, including the order in which the lines of code are executed.
* Explain at least two of the required components. For example, you could describe your conditional statement and what your loop does.

You can reference the example video here: https://drive.google.com/file/d/1LragCuaHO_eVMm1EDfFv7wqQcrC86NuW/view?usp=sharing

Upload the recorded video to your drive and paste the shareable link in the cell below.

Note: Ensure that we are able to view the video.

_Points:_ 10

```python
#https://drive.google.com/file/d/1348c1gO4M8BdsDV4rkMO73ArEwtLryIo/view?usp=drive_link
```

<!-- END QUESTION -->

### Submission

Before you submit your notebook, ensure that you've run all the cells sequentially to display images, graphs, and outputs correctly. Take the time to review your solutions and evaluate using the public test cases. Once you're satisfied, save your notebook by navigating to "File" in the Colab Notebook menu and selecting "Save" or using the keyboard shortcut (usually Ctrl + S or Cmd + S).

To submit, export your notebook as an .ipynb file. Ensure it's saved with the same name as the assignment. Then, upload this .ipynb file to Gradescope. Remember to submit the correct version.

Please save your work before exporting it, and if you encounter any technical issues or have questions about the submission process, reach out to the course staff for assistance.

---

To double-check your work for questions 1-9, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q1 results: All test cases passed!

q2 results: All test cases passed!

q3 results: All test cases passed!

q5 results: All test cases passed!

q6 results: All test cases passed!

q7 results: All test cases passed!

q8 results: All test cases passed!

q9 results: All test cases passed!
```

```python

```

```python

```

```python

```

Tests for q1 (points: 5)
```
>>> assert type(my_int) == str
```
```
>>> assert type(my_str) == int
```
```
>>> assert my_int == "Hi"
```
```
>>> assert my_str == 5
```


Tests for q2 (points: 5)
```
>>> assert calc_diag(3, 4) == 5
```


Tests for q3 (points: 10)
```
>>> assert is_divisible_by_3(9) == True
```
```
>>> assert is_divisible_by_3(-12) == True
```


Tests for q5 (points: 10)
```
>>> assert if_palindrome("racecar") == True
```
```
>>> assert if_palindrome("hello") == False
```


Tests for q6 (points: 5)
```
>>> assert len(sol) == 6
```
```
>>> assert sol == [5, 10, 15, 20, 25, 30]
>>>
```


Tests for q7 (points: 15)
```
>>> assert fibonacci(1) == [0]
```
```
>>> assert fibonacci(5) == [0, 1, 1, 2, 3]
```


Tests for q8 (points: 10)
```
>>> assert remove_vowels("Hello, World!") == "Hll, Wrld!"
```
```
>>> assert remove_vowels("Python Programming") == "Pythn Prgrmmng"
```


Tests for q9 (points: 10)
```
>>> assert sum(1, 5) == 9
>>>
```
```
>>> assert sum(-3, 3) == 0
```
